<?xml version="1.0" encoding="UTF-8"?>
<quiz>

<!-- ============================================================
     Q1: Competitive Equilibrium
     Supply: P = (Q + c) / d    Demand: P = (a - Q) / b
     Equilibrium: pStar = (a+c)/(b+d),  qStar = a - b*pStar
     ============================================================ -->
<question type="stack">
  <name><text>Microeconomics 1 - Competitive Equilibrium</text></name>

  <questiontext format="html">
    <text><![CDATA[
<p>The graph shows a demand curve and a supply curve for a competitive market.
<strong>Drag the orange point to the competitive equilibrium.</strong></p>

[[jsxgraph width="560px" height="460px"
  input-ref-xclick="xRef" input-ref-yclick="yRef"]]

  var a={#a#}, b={#b#}, c={#c#}, d={#d#};
  var pStar={#pStar#}, qStar={#qStar#};
  var xMax=Math.ceil(qStar*2.2/10)*10;
  var yMax=Math.ceil(Math.max(pStar*2.2, a/b)/10)*10;

  var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [0, yMax, xMax, 0],
    axis: true, grid: true,
    zoom: {enabled:false}, pan: {enabled:false},
    showNavigation: false,
    defaultAxes: {
      x: {name:'Q', withLabel:true, label:{position:'rt', offset:[0,-15]}},
      y: {name:'P', withLabel:true, label:{position:'rt', offset:[15,0]}}
    }
  });

  board.create('functiongraph',
    [function(q){return (a-q)/b;}, 0, xMax],
    {strokeColor:'black', strokeWidth:2,
     name:'Demand', withLabel:true,
     label:{position:'lft', offset:[10,10]}});

  board.create('functiongraph',
    [function(q){return (q+c)/d;}, 0, xMax],
    {strokeColor:'black', strokeWidth:2,
     name:'Supply', withLabel:true,
     label:{position:'lft', offset:[10,-15]}});

  var pt = board.create('point', [qStar*0.35, pStar*1.7], {
    color:'orange', size:6, name:'Your answer',
    label:{visible:false}, snapToGrid:false
  });

  function updateInputs() {
    document.getElementById(xRef).value = pt.X().toFixed(4);
    document.getElementById(yRef).value = pt.Y().toFixed(4);
  }
  pt.on('drag', updateInputs);
  pt.on('up',   updateInputs);
  updateInputs();

[[/jsxgraph]]

<div style="display:none;">
  [[input:xclick]] [[input:yclick]]
</div>
[[validation:xclick]] [[validation:yclick]]
    ]]></text>
  </questiontext>

  <generalfeedback format="html"><text></text></generalfeedback>
  <defaultgrade>1</defaultgrade>
  <penalty>0</penalty>
  <hidden>0</hidden>
  <idnumber></idnumber>
  <stackversion><text>2023010400</text></stackversion>

  <questionvariables>
    <text>
a: rand_with_step(80, 100, 1);
b: rand_with_step(10, 30, 1)/10;
c: rand_with_step(2, 8, 1);
d: rand_with_step(5, 20, 1)/10;
pStar: float((a + c) / (b + d));
qStar: float(a - b * pStar);
    </text>
  </questionvariables>

  <specificfeedback format="html">
    <text><![CDATA[[[feedback:prt1]]]]></text>
  </specificfeedback>

  <questionnote format="html">
    <text><![CDATA[a={#a#}, b={#b#}, c={#c#}, d={#d#} &rarr; Q*={#round(qStar,2)#}, P*={#round(pStar,2)#}]]></text>
  </questionnote>

  <questionsimplify>1</questionsimplify>
  <assumepositive>0</assumepositive>
  <assumereal>0</assumereal>

  <prtcorrect format="html">
    <text><![CDATA[<span style="color:green;font-weight:bold">Correct.</span>
    Equilibrium: Q* = {#round(qStar,2)#}, P* = {#round(pStar,2)#}.]]></text>
  </prtcorrect>
  <prtpartiallycorrect format="html">
    <text><![CDATA[<span style="color:orange">Partially correct.</span>]]></text>
  </prtpartiallycorrect>
  <prtincorrect format="html">
    <text><![CDATA[<span style="color:red;font-weight:bold">Incorrect.</span>
    The equilibrium was Q* = {#round(qStar,2)#}, P* = {#round(pStar,2)#}.]]></text>
  </prtincorrect>

  <multiplicationsign>dot</multiplicationsign>
  <sqrtsign>1</sqrtsign>
  <complexno>i</complexno>
  <inversetrig>cos-1</inversetrig>
  <logicsymbol>lang</logicsymbol>
  <matrixparens>[</matrixparens>
  <variantsselectionseed></variantsselectionseed>

  <input>
    <name>xclick</name>
    <type>numerical</type>
    <tans>qStar</tans>
    <boxsize>5</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint></syntaxhint>
    <syntaxattribute>0</syntaxattribute>
    <forbidwords></forbidwords>
    <allowwords></allowwords>
    <forbidfloat>0</forbidfloat>
    <requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype>
    <mustverify>0</mustverify>
    <showvalidation>0</showvalidation>
    <options></options>
  </input>

  <input>
    <name>yclick</name>
    <type>numerical</type>
    <tans>pStar</tans>
    <boxsize>5</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint></syntaxhint>
    <syntaxattribute>0</syntaxattribute>
    <forbidwords></forbidwords>
    <allowwords></allowwords>
    <forbidfloat>0</forbidfloat>
    <requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype>
    <mustverify>0</mustverify>
    <showvalidation>0</showvalidation>
    <options></options>
  </input>

  <prt>
    <name>prt1</name>
    <value>1</value>
    <autosimplify>1</autosimplify>
    <feedbackstyle>1</feedbackstyle>
    <feedbackvariables>
      <text>
qCorrect: abs(xclick - qStar) < 4;
pCorrect: abs(yclick - pStar) < 4;
      </text>
    </feedbackvariables>
    <node>
      <name>0</name>
      <description>Check Q coordinate</description>
      <answertest>AlgEquiv</answertest>
      <sans>qCorrect</sans>
      <tans>true</tans>
      <testoptions></testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>0</truescore>
      <truepenalty></truepenalty>
      <truenextnode>1</truenextnode>
      <trueanswernote>prt1-Q-T</trueanswernote>
      <truefeedback format="html"><text></text></truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0</falsescore>
      <falsepenalty></falsepenalty>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt1-Q-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[The quantity (x-axis) is not correct.]]></text>
      </falsefeedback>
    </node>
    <node>
      <name>1</name>
      <description>Check P coordinate</description>
      <answertest>AlgEquiv</answertest>
      <sans>pCorrect</sans>
      <tans>true</tans>
      <testoptions></testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>1</truescore>
      <truepenalty></truepenalty>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt1-P-T</trueanswernote>
      <truefeedback format="html"><text></text></truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0</falsescore>
      <falsepenalty></falsepenalty>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt1-P-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[The price (y-axis) is not correct.]]></text>
      </falsefeedback>
    </node>
  </prt>

</question>


<!-- ============================================================
     Q2: Unit Tax on Suppliers
     Original supply: P = (Q+c)/d
     Tax supply:      P = (Q+c)/d + t
     New equilibrium: qStar = (d*a - b*c - b*d*t)/(b+d)
                      pStar = (a - qStar)/b  [consumer price]
     ============================================================ -->
<question type="stack">
  <name><text>Microeconomics 2 - Unit Tax on Suppliers</text></name>

  <questiontext format="html">
    <text><![CDATA[
<p>The graph shows demand, the original supply, and the supply curve after a per-unit tax
of \(t = {#t#}\) is imposed on suppliers.
<strong>Drag the orange point to the new market equilibrium.</strong></p>

[[jsxgraph width="560px" height="460px"
  input-ref-xclick="xRef" input-ref-yclick="yRef"]]

  var a={#a#}, b={#b#}, c={#c#}, d={#d#}, t={#t#};
  var qOrig={#qOrig#}, pOrig={#pOrig#};
  var qStar={#qStar#}, pStar={#pStar#};
  var xMax=Math.ceil(Math.max(qOrig,qStar)*2.3/10)*10;
  var yMax=Math.ceil(Math.max(pStar*2.3, a/b)/10)*10;

  var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [0, yMax, xMax, 0],
    axis: true, grid: true,
    zoom: {enabled:false}, pan: {enabled:false},
    showNavigation: false,
    defaultAxes: {
      x: {name:'Q', withLabel:true, label:{position:'rt', offset:[0,-15]}},
      y: {name:'P', withLabel:true, label:{position:'rt', offset:[15,0]}}
    }
  });

  board.create('functiongraph',
    [function(q){return (a-q)/b;}, 0, xMax],
    {strokeColor:'black', strokeWidth:2,
     name:'Demand', withLabel:true,
     label:{position:'lft', offset:[10,10]}});

  board.create('functiongraph',
    [function(q){return (q+c)/d;}, 0, xMax],
    {strokeColor:'black', strokeWidth:2, dash:2,
     name:'Supply (original)', withLabel:true,
     label:{position:'lft', offset:[10,-10]}});

  board.create('functiongraph',
    [function(q){return (q+c)/d + t;}, 0, xMax],
    {strokeColor:'black', strokeWidth:2,
     name:'Supply (after tax)', withLabel:true,
     label:{position:'lft', offset:[10,-25]}});

  var pt = board.create('point', [qStar*0.35, pStar*1.7], {
    color:'orange', size:6, name:'Your answer',
    label:{visible:false}, snapToGrid:false
  });

  function updateInputs() {
    document.getElementById(xRef).value = pt.X().toFixed(4);
    document.getElementById(yRef).value = pt.Y().toFixed(4);
  }
  pt.on('drag', updateInputs);
  pt.on('up',   updateInputs);
  updateInputs();

[[/jsxgraph]]

<div style="display:none;">
  [[input:xclick]] [[input:yclick]]
</div>
[[validation:xclick]] [[validation:yclick]]
    ]]></text>
  </questiontext>

  <generalfeedback format="html"><text></text></generalfeedback>
  <defaultgrade>1</defaultgrade>
  <penalty>0</penalty>
  <hidden>0</hidden>
  <idnumber></idnumber>
  <stackversion><text>2023010400</text></stackversion>

  <questionvariables>
    <text>
a: rand_with_step(80, 100, 1);
b: rand_with_step(10, 30, 1)/10;
c: rand_with_step(2, 8, 1);
d: rand_with_step(5, 20, 1)/10;
t: rand_with_step(30, 80, 1)/10;
pOrig: float((a + c) / (b + d));
qOrig: float(a - b * pOrig);
qStar: float((d*a - b*c - b*d*t) / (b + d));
pStar: float((a - qStar) / b);
    </text>
  </questionvariables>

  <specificfeedback format="html">
    <text><![CDATA[[[feedback:prt1]]]]></text>
  </specificfeedback>

  <questionnote format="html">
    <text><![CDATA[a={#a#}, b={#b#}, c={#c#}, d={#d#}, t={#t#} &rarr; Q*={#round(qStar,2)#}, P*={#round(pStar,2)#}]]></text>
  </questionnote>

  <questionsimplify>1</questionsimplify>
  <assumepositive>0</assumepositive>
  <assumereal>0</assumereal>

  <prtcorrect format="html">
    <text><![CDATA[<span style="color:green;font-weight:bold">Correct.</span>
    New equilibrium: Q* = {#round(qStar,2)#}, P* = {#round(pStar,2)#} (consumer price).]]></text>
  </prtcorrect>
  <prtpartiallycorrect format="html">
    <text><![CDATA[<span style="color:orange">Partially correct.</span>]]></text>
  </prtpartiallycorrect>
  <prtincorrect format="html">
    <text><![CDATA[<span style="color:red;font-weight:bold">Incorrect.</span>
    The new equilibrium was Q* = {#round(qStar,2)#}, P* = {#round(pStar,2)#}.]]></text>
  </prtincorrect>

  <multiplicationsign>dot</multiplicationsign>
  <sqrtsign>1</sqrtsign>
  <complexno>i</complexno>
  <inversetrig>cos-1</inversetrig>
  <logicsymbol>lang</logicsymbol>
  <matrixparens>[</matrixparens>
  <variantsselectionseed></variantsselectionseed>

  <input>
    <name>xclick</name>
    <type>numerical</type>
    <tans>qStar</tans>
    <boxsize>5</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint></syntaxhint>
    <syntaxattribute>0</syntaxattribute>
    <forbidwords></forbidwords>
    <allowwords></allowwords>
    <forbidfloat>0</forbidfloat>
    <requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype>
    <mustverify>0</mustverify>
    <showvalidation>0</showvalidation>
    <options></options>
  </input>

  <input>
    <name>yclick</name>
    <type>numerical</type>
    <tans>pStar</tans>
    <boxsize>5</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint></syntaxhint>
    <syntaxattribute>0</syntaxattribute>
    <forbidwords></forbidwords>
    <allowwords></allowwords>
    <forbidfloat>0</forbidfloat>
    <requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype>
    <mustverify>0</mustverify>
    <showvalidation>0</showvalidation>
    <options></options>
  </input>

  <prt>
    <name>prt1</name>
    <value>1</value>
    <autosimplify>1</autosimplify>
    <feedbackstyle>1</feedbackstyle>
    <feedbackvariables>
      <text>
qCorrect: abs(xclick - qStar) < 4;
pCorrect: abs(yclick - pStar) < 4;
      </text>
    </feedbackvariables>
    <node>
      <name>0</name>
      <description>Check Q</description>
      <answertest>AlgEquiv</answertest>
      <sans>qCorrect</sans>
      <tans>true</tans>
      <testoptions></testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>0</truescore>
      <truepenalty></truepenalty>
      <truenextnode>1</truenextnode>
      <trueanswernote>prt1-Q-T</trueanswernote>
      <truefeedback format="html"><text></text></truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0</falsescore>
      <falsepenalty></falsepenalty>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt1-Q-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[The quantity (x-axis) is not correct.]]></text>
      </falsefeedback>
    </node>
    <node>
      <name>1</name>
      <description>Check P</description>
      <answertest>AlgEquiv</answertest>
      <sans>pCorrect</sans>
      <tans>true</tans>
      <testoptions></testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>1</truescore>
      <truepenalty></truepenalty>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt1-P-T</trueanswernote>
      <truefeedback format="html"><text></text></truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0</falsescore>
      <falsepenalty></falsepenalty>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt1-P-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[The price (y-axis) is not correct.]]></text>
      </falsefeedback>
    </node>
  </prt>

</question>


<!-- ============================================================
     Q3: Constrained Utility Maximization
     U = x*y,  budget: px*x + py*y = I
     Optimum:  xStar = I/(2*px),  yStar = I/(2*py)
     Three ICs: uLow = uStar*0.45, uMid = uStar, uHigh = uStar*1.65
     ============================================================ -->
<question type="stack">
  <name><text>Microeconomics 3 - Constrained Utility Maximization</text></name>

  <questiontext format="html">
    <text><![CDATA[
<p>The graph shows three indifference curves (\(U = x \cdot y\)) and a budget line
(\(P_x = {#px#}\), \(P_y = {#py#}\), \(I = {#I#}\)).
<strong>Drag the orange point to the constrained utility-maximizing bundle.</strong></p>

[[jsxgraph width="560px" height="460px"
  input-ref-xclick="xRef" input-ref-yclick="yRef"]]

  var px={#px#}, py={#py#}, Inco={#I#};
  var xStar={#xStar#}, yStar={#yStar#};
  var uStar={#uStar#};
  var uLow=uStar*0.45, uHigh=uStar*1.65;

  var xBudget=Inco/px, yBudget=Inco/py;
  var xMax=xBudget*1.3, yMax=yBudget*1.3;

  var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [0, yMax, xMax, 0],
    axis: true, grid: true,
    zoom: {enabled:false}, pan: {enabled:false},
    showNavigation: false,
    defaultAxes: {
      x: {name:'x', withLabel:true, label:{position:'rt', offset:[0,-15]}},
      y: {name:'y', withLabel:true, label:{position:'rt', offset:[15,0]}}
    }
  });

  // Indifference curves  y = U / x
  board.create('functiongraph',
    [function(x){return uLow/x;}, 0.3, xMax],
    {strokeColor:'black', strokeWidth:2,
     name:'IC\u2081', withLabel:true,
     label:{position:'lft', offset:[5,10]}});

  board.create('functiongraph',
    [function(x){return uStar/x;}, 0.3, xMax],
    {strokeColor:'black', strokeWidth:2,
     name:'IC\u2082', withLabel:true,
     label:{position:'lft', offset:[5,10]}});

  board.create('functiongraph',
    [function(x){return uHigh/x;}, 0.3, xMax],
    {strokeColor:'black', strokeWidth:2,
     name:'IC\u2083', withLabel:true,
     label:{position:'lft', offset:[5,10]}});

  // Budget line  y = (I - px*x) / py
  board.create('functiongraph',
    [function(x){return (Inco - px*x)/py;}, 0, xBudget],
    {strokeColor:'black', strokeWidth:2, dash:2,
     name:'Budget', withLabel:true,
     label:{position:'lft', offset:[5,-15]}});

  var pt = board.create('point', [xStar*0.3, yStar*1.8], {
    color:'orange', size:6, name:'Your answer',
    label:{visible:false}, snapToGrid:false
  });

  function updateInputs() {
    document.getElementById(xRef).value = pt.X().toFixed(4);
    document.getElementById(yRef).value = pt.Y().toFixed(4);
  }
  pt.on('drag', updateInputs);
  pt.on('up',   updateInputs);
  updateInputs();

[[/jsxgraph]]

<div style="display:none;">
  [[input:xclick]] [[input:yclick]]
</div>
[[validation:xclick]] [[validation:yclick]]
    ]]></text>
  </questiontext>

  <generalfeedback format="html"><text></text></generalfeedback>
  <defaultgrade>1</defaultgrade>
  <penalty>0</penalty>
  <hidden>0</hidden>
  <idnumber></idnumber>
  <stackversion><text>2023010400</text></stackversion>

  <questionvariables>
    <text>
px:   rand_with_step(10, 30, 1)/10;
py:   rand_with_step(10, 30, 1)/10;
I:    rand_with_step(20, 40, 1);
xStar: float(I / (2 * px));
yStar: float(I / (2 * py));
uStar: float(xStar * yStar);
tolX: float(max(1, xStar * 0.12));
tolY: float(max(1, yStar * 0.12));
    </text>
  </questionvariables>

  <specificfeedback format="html">
    <text><![CDATA[[[feedback:prt1]]]]></text>
  </specificfeedback>

  <questionnote format="html">
    <text><![CDATA[px={#px#}, py={#py#}, I={#I#} &rarr; x*={#round(xStar,2)#}, y*={#round(yStar,2)#}]]></text>
  </questionnote>

  <questionsimplify>1</questionsimplify>
  <assumepositive>0</assumepositive>
  <assumereal>0</assumereal>

  <prtcorrect format="html">
    <text><![CDATA[<span style="color:green;font-weight:bold">Correct.</span>
    Optimal bundle: x* = {#round(xStar,2)#}, y* = {#round(yStar,2)#}.]]></text>
  </prtcorrect>
  <prtpartiallycorrect format="html">
    <text><![CDATA[<span style="color:orange">Partially correct.</span>]]></text>
  </prtpartiallycorrect>
  <prtincorrect format="html">
    <text><![CDATA[<span style="color:red;font-weight:bold">Incorrect.</span>
    The optimal bundle was x* = {#round(xStar,2)#}, y* = {#round(yStar,2)#}.]]></text>
  </prtincorrect>

  <multiplicationsign>dot</multiplicationsign>
  <sqrtsign>1</sqrtsign>
  <complexno>i</complexno>
  <inversetrig>cos-1</inversetrig>
  <logicsymbol>lang</logicsymbol>
  <matrixparens>[</matrixparens>
  <variantsselectionseed></variantsselectionseed>

  <input>
    <name>xclick</name>
    <type>numerical</type>
    <tans>xStar</tans>
    <boxsize>5</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint></syntaxhint>
    <syntaxattribute>0</syntaxattribute>
    <forbidwords></forbidwords>
    <allowwords></allowwords>
    <forbidfloat>0</forbidfloat>
    <requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype>
    <mustverify>0</mustverify>
    <showvalidation>0</showvalidation>
    <options></options>
  </input>

  <input>
    <name>yclick</name>
    <type>numerical</type>
    <tans>yStar</tans>
    <boxsize>5</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint></syntaxhint>
    <syntaxattribute>0</syntaxattribute>
    <forbidwords></forbidwords>
    <allowwords></allowwords>
    <forbidfloat>0</forbidfloat>
    <requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype>
    <mustverify>0</mustverify>
    <showvalidation>0</showvalidation>
    <options></options>
  </input>

  <prt>
    <name>prt1</name>
    <value>1</value>
    <autosimplify>1</autosimplify>
    <feedbackstyle>1</feedbackstyle>
    <feedbackvariables>
      <text>
xCorrect: abs(xclick - xStar) < tolX;
yCorrect: abs(yclick - yStar) < tolY;
      </text>
    </feedbackvariables>
    <node>
      <name>0</name>
      <description>Check x coordinate</description>
      <answertest>AlgEquiv</answertest>
      <sans>xCorrect</sans>
      <tans>true</tans>
      <testoptions></testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>0</truescore>
      <truepenalty></truepenalty>
      <truenextnode>1</truenextnode>
      <trueanswernote>prt1-x-T</trueanswernote>
      <truefeedback format="html"><text></text></truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0</falsescore>
      <falsepenalty></falsepenalty>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt1-x-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[The quantity of good x is not correct.]]></text>
      </falsefeedback>
    </node>
    <node>
      <name>1</name>
      <description>Check y coordinate</description>
      <answertest>AlgEquiv</answertest>
      <sans>yCorrect</sans>
      <tans>true</tans>
      <testoptions></testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>1</truescore>
      <truepenalty></truepenalty>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt1-y-T</trueanswernote>
      <truefeedback format="html"><text></text></truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0</falsescore>
      <falsepenalty></falsepenalty>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt1-y-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[The quantity of good y is not correct.]]></text>
      </falsefeedback>
    </node>
  </prt>

</question>


<!-- ============================================================
     Q4: Negative Production Externality
     Demand: P = a - b*Q
     PMC:    P = c + d*Q
     SMC:    P = c + (d+e)*Q   (MED = e*Q)
     Social optimum: qStar = (a-c)/(b+d+e),  pStar = a - b*qStar
     ============================================================ -->
<question type="stack">
  <name><text>Microeconomics 4 - Negative Production Externality</text></name>

  <questiontext format="html">
    <text><![CDATA[
<p>The graph shows demand, private marginal cost (PMC), and social marginal cost (SMC)
for a market with a negative production externality.
<strong>Drag the orange point to the socially efficient equilibrium.</strong></p>

[[jsxgraph width="560px" height="460px"
  input-ref-xclick="xRef" input-ref-yclick="yRef"]]

  var a={#a#}, b={#b#}, c={#c#}, d={#d#}, e={#e#};
  var qStar={#qStar#}, pStar={#pStar#};
  var qMarket={#qMarket#};
  var xMax=Math.ceil(qMarket*2.2/10)*10;
  var yMax=Math.ceil(Math.max(pStar*2.2, a*0.9)/10)*10;

  var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [0, yMax, xMax, 0],
    axis: true, grid: true,
    zoom: {enabled:false}, pan: {enabled:false},
    showNavigation: false,
    defaultAxes: {
      x: {name:'Q', withLabel:true, label:{position:'rt', offset:[0,-15]}},
      y: {name:'P', withLabel:true, label:{position:'rt', offset:[15,0]}}
    }
  });

  board.create('functiongraph',
    [function(q){return a - b*q;}, 0, xMax],
    {strokeColor:'black', strokeWidth:2,
     name:'Demand', withLabel:true,
     label:{position:'lft', offset:[5,10]}});

  board.create('functiongraph',
    [function(q){return c + d*q;}, 0, xMax],
    {strokeColor:'black', strokeWidth:2,
     name:'PMC', withLabel:true,
     label:{position:'lft', offset:[5,-10]}});

  board.create('functiongraph',
    [function(q){return c + (d+e)*q;}, 0, xMax],
    {strokeColor:'black', strokeWidth:2, dash:2,
     name:'SMC', withLabel:true,
     label:{position:'lft', offset:[5,-25]}});

  var pt = board.create('point', [qStar*0.35, pStar*1.7], {
    color:'orange', size:6, name:'Your answer',
    label:{visible:false}, snapToGrid:false
  });

  function updateInputs() {
    document.getElementById(xRef).value = pt.X().toFixed(4);
    document.getElementById(yRef).value = pt.Y().toFixed(4);
  }
  pt.on('drag', updateInputs);
  pt.on('up',   updateInputs);
  updateInputs();

[[/jsxgraph]]

<div style="display:none;">
  [[input:xclick]] [[input:yclick]]
</div>
[[validation:xclick]] [[validation:yclick]]
    ]]></text>
  </questiontext>

  <generalfeedback format="html"><text></text></generalfeedback>
  <defaultgrade>1</defaultgrade>
  <penalty>0</penalty>
  <hidden>0</hidden>
  <idnumber></idnumber>
  <stackversion><text>2023010400</text></stackversion>

  <questionvariables>
    <text>
a: rand_with_step(70, 100, 1);
b: rand_with_step(8, 20, 1)/10;
c: rand_with_step(5, 20, 1);
d: rand_with_step(5, 15, 1)/10;
e: rand_with_step(5, 20, 1)/10;
qStar:   float((a - c) / (b + d + e));
pStar:   float(a - b * qStar);
qMarket: float((a - c) / (b + d));
    </text>
  </questionvariables>

  <specificfeedback format="html">
    <text><![CDATA[[[feedback:prt1]]]]></text>
  </specificfeedback>

  <questionnote format="html">
    <text><![CDATA[a={#a#}, b={#b#}, c={#c#}, d={#d#}, e={#e#} &rarr; Q*={#round(qStar,2)#}, P*={#round(pStar,2)#}]]></text>
  </questionnote>

  <questionsimplify>1</questionsimplify>
  <assumepositive>0</assumepositive>
  <assumereal>0</assumereal>

  <prtcorrect format="html">
    <text><![CDATA[<span style="color:green;font-weight:bold">Correct.</span>
    Social optimum: Q* = {#round(qStar,2)#}, P* = {#round(pStar,2)#}.]]></text>
  </prtcorrect>
  <prtpartiallycorrect format="html">
    <text><![CDATA[<span style="color:orange">Partially correct.</span>]]></text>
  </prtpartiallycorrect>
  <prtincorrect format="html">
    <text><![CDATA[<span style="color:red;font-weight:bold">Incorrect.</span>
    The social optimum was Q* = {#round(qStar,2)#}, P* = {#round(pStar,2)#}.]]></text>
  </prtincorrect>

  <multiplicationsign>dot</multiplicationsign>
  <sqrtsign>1</sqrtsign>
  <complexno>i</complexno>
  <inversetrig>cos-1</inversetrig>
  <logicsymbol>lang</logicsymbol>
  <matrixparens>[</matrixparens>
  <variantsselectionseed></variantsselectionseed>

  <input>
    <name>xclick</name>
    <type>numerical</type>
    <tans>qStar</tans>
    <boxsize>5</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint></syntaxhint>
    <syntaxattribute>0</syntaxattribute>
    <forbidwords></forbidwords>
    <allowwords></allowwords>
    <forbidfloat>0</forbidfloat>
    <requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype>
    <mustverify>0</mustverify>
    <showvalidation>0</showvalidation>
    <options></options>
  </input>

  <input>
    <name>yclick</name>
    <type>numerical</type>
    <tans>pStar</tans>
    <boxsize>5</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint></syntaxhint>
    <syntaxattribute>0</syntaxattribute>
    <forbidwords></forbidwords>
    <allowwords></allowwords>
    <forbidfloat>0</forbidfloat>
    <requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype>
    <mustverify>0</mustverify>
    <showvalidation>0</showvalidation>
    <options></options>
  </input>

  <prt>
    <name>prt1</name>
    <value>1</value>
    <autosimplify>1</autosimplify>
    <feedbackstyle>1</feedbackstyle>
    <feedbackvariables>
      <text>
qCorrect: abs(xclick - qStar) < 4;
pCorrect: abs(yclick - pStar) < 4;
      </text>
    </feedbackvariables>
    <node>
      <name>0</name>
      <description>Check Q</description>
      <answertest>AlgEquiv</answertest>
      <sans>qCorrect</sans>
      <tans>true</tans>
      <testoptions></testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>0</truescore>
      <truepenalty></truepenalty>
      <truenextnode>1</truenextnode>
      <trueanswernote>prt1-Q-T</trueanswernote>
      <truefeedback format="html"><text></text></truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0</falsescore>
      <falsepenalty></falsepenalty>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt1-Q-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[The quantity (x-axis) is not correct.]]></text>
      </falsefeedback>
    </node>
    <node>
      <name>1</name>
      <description>Check P</description>
      <answertest>AlgEquiv</answertest>
      <sans>pCorrect</sans>
      <tans>true</tans>
      <testoptions></testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>1</truescore>
      <truepenalty></truepenalty>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt1-P-T</trueanswernote>
      <truefeedback format="html"><text></text></truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0</falsescore>
      <falsepenalty></falsepenalty>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt1-P-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[The price (y-axis) is not correct.]]></text>
      </falsefeedback>
    </node>
  </prt>

</question>

</quiz>
